```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

<a id='verytop'></a>

# Lab 2: Linear Transformations and Geometry 

## Due Date: Thurs, Feb 5th 11:59 PM on Gradescope


### Detailed Submission Instructions Are Provided at the end of this Notebook


## Collaboration Policy

A key step in learning and retention is **creating solutions on your own.**   Below are examples of acceptable vs unacceptable use of resources and collaboration when doing lab assignments in CSCI 2820.


The following would be some **examples of cheating** when working on HW assignments in CSCI 2820.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given lab problem or one similar is cheating.  However, consulting notes from the class videos, and web pages that explain the material taught in class but do NOT show a solution to the lab problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course videos, the course textbooks, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the lab before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a lab problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a lab solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit office Hours and/or post questions on Piazza!



## Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.

### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1abcef | 8 | autograded
Question 1d | 2 | manual
Question 2ab | 4 | autograded
Question 3a | 2 | autograded
Question 3b | 2 | manual (1) , autograded (1)
Question 4a | 2 | autograded
Question 4b | 2 | manual (1) , autograded (1)
Question 5a | 2 | autograded
Question 5b | 2 | manual (1) , autograded (1)
Question 5c | 3 | manual (2), autograded (1)
Question 5d | 2 | autograded
Question 5e | 2 | manual (1) , autograded (1)
Question 6a | 2 | autograded
Question 6b | 2 | manual
Question 7a | 2 | autograded
Question 7b | 2 | manual
Question 8a | 3 | autograded
Question 8b | 2 | manual
Question 9a | 2 | autograded
Question 9b | 2 | manual
|Total | 50


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


```python
# import SymPy with alias sp
import sympy as sp
import numpy as np

# Function needed to run in-notebook tests
import hashlib

def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()


def get_array_hash_normalized(arr):
    """Hash numpy array ignoring dtype differences (values only)."""
    arr = np.asarray(arr, dtype=float)  # normalize to float
    return hashlib.md5(arr.tobytes() + str(arr.shape).encode()).hexdigest()
```

### NumPy module:

𝐍𝐮𝐦𝐏𝐲  (short for "Numeric Python") is an open source Python module for scientific computing. NumPy supports large, multidimensional arrays and matrices. It also supports a large collection of mathematical functions not found in Python's standard math library.

It's custom in data science to import Numpy  with the alias $\texttt{np}$.  We can then ***access any mathematical function in the Numpy library by prepending function names by $\texttt{np}$***

For a more in-depth description of methods and attributes built into Numpy see here:  https://numpy.org/doc/stable/user/quickstart.html   
and  here:  https://numpy.org/doc/stable/reference/routines.math.html

*Hint: If you press `shift+tab` while next to the function call, the documentation for that function will appear.*

**Vectors in NumPy:**

To create the vector $$\mathbf{v} = \left[\begin{array}{r} -1 \\ 3 \\ 2 \end{array} \right]$$  in NumPy we can use the following syntax:

```python
v = np.array([[-1],[3], [2]])
v
```

```python
# To create a 3D vector of all zeros:
np.zeros((3,1))
```

**Matrices in NumPy:**

To create the matrix $$
\begin{equation}
A = \left[\begin{array}{rrr} 1 & 1 & 0 \\ 1 & 0 & 2 \\ 3 & 1 & -1 \end{array}\right]
\end{equation}
$$ in NumPy we can use the following syntax:

```python
A = np.array([[1,1,0],[1,0,2],[3,1,-1]])
A
```

**Accessing elements in vectors and matrices in NumPy**

Accessing elements in a NumPy array is done through indexing.  
You provide indices for each dimension, separated by commas, within a single set of square brackets. The order of indices typically corresponds to (row, column) for 2D arrays.  NumPy arrays are 0-indexed, meaning the first element is at index 0. 


```python
# To access the element in the 2nd row and 3rd column in A:
A[1,2]
```

```python
# To access the 2nd entry in the vector v:

v[1,0]
```

**Matrix-Vector or Matrix-Matrix Multiplication in NumPy**

## WARNING:   The * operator does NOT compute Matrix Multiplication - you must use @ instead

In NumPy, the `*` operator does **not** perform matrix multiplication because NumPy was designed to treat arrays primarily as collections of numbers, not as linear-algebra objects by default. To avoid ambiguity and to support fast, element-wise numerical computation, `*` is defined as **element-wise multiplication**, not matrix multiplication.

Instead, matrix multiplication is done explicitly using `@` (or `np.matmul`), which makes the programmer’s intent clear.


```python
#Multiply A and v by hand and then run this cell to check your answer.

A@v
```

```python
#Notice, if you use *, this returns:

print(A) 
print(v)
A*v
```


When you use `*`, NumPy applies **broadcasting** if the arrays have different shapes. Broadcasting is a rule-based system that allows NumPy to perform element-wise operations on arrays of different shapes by *virtually* expanding the smaller array (without copying data) so their shapes are compatible. Dimensions of size 1 can be stretched to match the corresponding dimension of the other array.

Example:

* Multiplying a matrix by a vector with `*` multiplies each row (or column) element-wise, not via dot products.
* A scalar times an array works because the scalar is broadcast to every entry.

In short **DON'T USE `*` WHEN DOING LINEAR ALGEBRA MULTIPLICATION IN NUMPY**



### Linear Transformations

<img src="img/mapping.png" width = 800> 

### Example 1:  Transformation from $\mathbb{R}^2$ to $\mathbb{R}^3$

A Python function that accepts one vector as an argument, and provides one vector in return, is a perfect model of a mapping.  As an example, let's write a function that accepts a vector from $\mathbb{R}^2$ and returns a vector in $\mathbb{R}^3$.

```python
# Run this cell:
def T(V):
    #Create a 3D vector with all zeros
    W = np.zeros((3,1))
    #Assign the first element in the vector W to be twice the first element in vector V
    W[0,0] = 2*V[0,0]
    # Assign the 3rd element in the vector W to be the 2nd element in the vector V
    W[2,0] = V[1,0]
    return W


```

```python
# Evaluate T(V) for a specific vector V:

V = np.array([[3],[5]])
W = T(V)

print(V,'\n\n',W)
```

We can describe this mapping $T:\mathbb{R}^2 \to \mathbb{R}^3$ with the following symbolic formula.

$$
\begin{equation}
T \left(\left[\begin{array}{r} v_1 \\ v_2 \end{array} \right]\right) = 
\left[\begin{array}{r} 2v_1 \\ 0  \\ v_2 \end{array} \right]
\end{equation}
$$



### Linearity

A map $T$ is said to be *linear* if the following two properties are true for any scalar $k$, and any input vectors $\mathbf{u}$ and $\mathbf{v}$.  

- $T(k\mathbf{v})$ = $kT(\mathbf{v})$
- $T(\mathbf{u}+ \mathbf{v})$ = $T(\mathbf{u}) + T(\mathbf{v})$

A linear map is referred to as a **linear transformation**.

The map $T$ in **Example 1** is a linear transformation.   We can prove this as follows: 
$$
\begin{eqnarray*}
T(k\mathbf{v})  = T \left(\left [\begin{array}{r} kv_1 \\ kv_2 \end{array} \right]\right) & = & 
\left[\begin{array}{r} 2kv_1 \\ 0  \\ kv_2 \end{array} \right] =
k\left[\begin{array}{r} 2v_1 \\ 0  \\ v_2 \end{array} \right] = kT(\mathbf{v})\\ 
\\
T(\mathbf{u+v}) = T \left(\left [\begin{array}{r} u_1 + v_1 \\ u_2 + v_2 \end{array} \right]\right) & = & 
\left[\begin{array}{r} 2u_1 \\ 0  \\ u_2 \end{array} \right] +
\left[\begin{array}{r} 2v_1 \\ 0  \\ v_2 \end{array} \right] = T(\mathbf{u}) + T(\mathbf{v})
\end{eqnarray*}
$$

If we choose specific vectors $\mathbf{u}$ and $\mathbf{v}$, and a specific scalar $k$, we can compute the left and right sides of these equations with syntax that matches the notation :


```python
U = np.array([[-1],[3]])
V = np.array([[3],[5]])
k = 7
print(T(k*V),'\n')
print(k*T(V),'\n')
print(T(U+V),'\n')
print(T(U)+T(V))
```

Showing these equations hold for a *specific choice* of $\mathbf{u}$, $\mathbf{v}$, and $k$ does not prove that the transformation is linear.  Even if we were to run this code inside a loop and see that the equations held for a million random combinations of $\mathbf{u}$ $\mathbf{v}$, and $k$, that would still not *prove* that the transformation is linear.  We need to use the mathematical proof given above (in LaTeX) that holds for any generic 2 vectors.    It would however provide numerical evidence to support the claim.  

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴  Question 1a (1 pt) ###

Let $T(\mathbf{v})$ be defined as above.  Find the vector $T(\mathbf{v})$ where

$$
\begin{equation}
\mathbf{v} = \left[\begin{array}{r} 1 \\ 3 \end{array} \right]
\end{equation}
$$

Since $T$ outputs vectors in $\mathbb{R}^3$ make sure your output is a 3x1 Numpy Array (not a 1x3 array). 

```python
V = ...

T_of_V = ...

#print T(V)
print(T_of_V)
```

```python
grader.check("q1a")
```

## 🔴   Question 1b (1 pt) ###

Find the vector $\mathbf{u}$ so that 

$$
\begin{equation}
T(\mathbf{u}) = \left[\begin{array}{r} 10 \\ 0 \\ -1 \end{array} \right]
\end{equation}
$$

(Do this by hand and enter in your answer below)


```python
U = ...


print(U)
```

```python
grader.check("q1b")
```

## 🔴   Question 1c (2 pts) ###

Write a Python function that implements the transformation $N:\mathbb{R}^3\to\mathbb{R}^2$, given by the following rule.  

$$
\begin{equation}
N \left(\left[\begin{array}{r} v_1 \\ v_2 \\ v_3 \end{array} \right]\right) = 
\left[\begin{array}{c} 8v_2 \\  v_1 + v_3 + 3 \end{array} \right]
\end{equation}
$$


```python
def N(V):
    ...
    
    return W
```

```python
grader.check("q1c")
```

<!-- BEGIN QUESTION -->

## 🔴   Question 1d (2 pts) ###

Is the transformation $N$ a linear transformation?  If so, prove it using the definition of linear transformations.  If not, give a counterexample.

**Type your work answering Question 1d** in this cell.   If your answer is yes, type your proof here.  If your answer is no, give a countexample in this cell (use LaTeX for any math notation).  Do not add any additional cells to this part.  

```python
# Optional cell for any code you want to use when exploring this question 
#(this cell will not be graded, so make sure your full answer is in the Markdown cell above)


```

<!-- END QUESTION -->

### Example 2:  Transformation defined as matrix multiplication

Let $L:\mathbb{R}^3 \to \mathbb{R}^3$ be defined by $L(\mathbf{x})= A\mathbf{x}$ where $A$ is the following matrix.

$$
\begin{equation}
A = \left[\begin{array}{rrr} 1 & 1 & 0 \\ 1 & 0 & 2 \\ 3 & 1 & -1 \end{array}\right]
\end{equation}
$$

Since matrix multiplication is linear, a transformation defined by matrix multiplication is linear as well.  

We can see this as follows:  If $\mathbf{x}$ and $\mathbf{y}$ are any vectors in $\mathbb{R}^3$, and $k$ is any scalar, then $$L(\mathbf{x+y}) = A(\mathbf{x+y}) = A\mathbf{x} + A\mathbf{y} = L(\mathbf{x}) + L(\mathbf{y})$$ and $$L(k\mathbf{x}) = A(k\mathbf{x}) = k(A\mathbf{x}) = kL(\mathbf{x})$$.  

The matrix $A$ is called the **standard matrix for the transformation**.  

Let's define this linear transformation in terms of matrix multiplication:



```python
def L(V):
    A = np.array([[1,1,0],[1,0,2],[3,1,-1]])
    W = A@V
    return W


```

```python
# Evaluate L(V) for a specific V

V = np.array([[1],[3],[5]])
W = L(V)

#Print out both V and W:
print(V,'\n\n',W)
```

## 🔴   Question 1e (2 pts) ###

What is the standard matrix for the transformation $T$, defined in Example 1 above?  (We have recopied the code defining T in the cell below for your convenience). 

```python

T_stand_matrix = ...


#Quick check that your answer is correct - check you get the same answer when applying the transformation to a specific v:

v = np.array([[5], [-1]])

print(T_stand_matrix@v)

print(T(v))

# Recall the definition of T(V) above:
#def T(V):
#    #Create a 3D vector with all zeros
#    W = np.zeros((3,1))
#    #Assign the first element in the vector W to be twice the first element in vector V
#    W[0,0] = 2*V[0,0]
#    # Assign the 3rd element in the vector W to be the 2nd element in the vector V
#    W[2,0] = V[1,0]
#    return W

```

```python
grader.check("q1e")
```

### Compositions

If two transformations are to be applied, one after the other, it is common to combine them to form what is known as a **composition**.  This idea only makes sense when the output space of the first transformation matches the input space of the second.  In symbols, if $T:\mathbb{U}\to\mathbb{V}$ and $L:\mathbb{V}\to\mathbb{W}$, then we could consider the composition $L\circ T:\mathbb{U}\to\mathbb{W}$.  The notation $L\circ T$ represents the composition formed by applying the map $T$, followed by the map $L$.

As an example, if we take $T:\mathbb{R}^2 \to \mathbb{R}^3$ and $L:\mathbb{R}^3 \to \mathbb{R}^3$ to be the transformations defined in **Example 1** and **Example 2**, we could define $L\circ T:\mathbb{R}^2 \to \mathbb{R}^3$.



## 🔴  Question 1f (2 pts) ###

What is the standard matrix of the composition $(L \circ T)$? 

```python

L_T_stand_matrix = ...


L_T_stand_matrix

#Check your answer by comparing these two - they should be the same:
u = np.array([[3],[5]])

print(L_T_stand_matrix@u)


composition_output = L(T(u))
print(composition_output)


```

```python
grader.check("q1f")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴   Question 2a (3 pts) 

 Consider the two liear transformations, $S$ and $R$, defined below.  



$$
\begin{equation}
S \left(\left[\begin{array}{r} v_1 \\ v_2 \\ v_3 \end{array} \right]\right) = 
\left[\begin{array}{c}   v_1 + v_2 \\  3v_3 \end{array} \right]
\end{equation}
$$

$$
\begin{equation}
R \left(\left[\begin{array}{r} v_1 \\ v_2  \end{array} \right]\right) = 
\left[\begin{array}{rr} 3 &  0 \\ -1 & 1 \end{array}\right]
\left[\begin{array}{c}   v_1 \\ v_2 \end{array} \right]
\end{equation}
$$


i). Give the standard matrix for the transformation S

ii). Write functions $S$ and $R$ that implement each transformation, using matrix multiplication.  


iii). Write a function that implements the composition $R\circ S$, by calling the functions S and R you have defined.



```python
S_stand_matrix = np.array([[1, 1, 0], [0, 0, 3]])



def S(V):
    ...
    return W

def R(V):
    ...
    return W

def R_composed_with_S(V):
    ...
    return W


```

```python
grader.check("2a")
```

## 🔴  Question 2b (1 pt) ###

**Multiple Choice:**
Is it possible to form the composition $S \circ R$? 

Assign the the variable  `ans_2b` to the corresponding letter (as a capital letter string `'A'`,`'B'`, or `'C'`) to  below to indicate your answer.  Replace the ellipses (`...`) with your answer. 

A).  Yes

B).  No


```python
ans_2b = ...


```

```python
grader.check("q2b")
```

## Transformations in a Plane

As discussed in the class lesson videos, any linear transformation $T:\mathbb{R}^2\to\mathbb{R}^2$ can be represented as the multiplication of a $2\times 2$ matrix and a coordinate vector.  For now, the coordinates are with respect to the standard basis $\{e_1,e_2\}$ as they were before.  The columns of the matrix are then the images of the the basis vectors.

$$
\begin{equation}
A = \left[ \begin{array}{c|c} T(e_1) \hspace{5mm}  T(e_2) \end{array} \right]
\end{equation}
$$

### Example 3:  Horizontal Stretch

As a first example let us consider the transformation defined by the following images.

$$
\begin{equation}
T\left(\left[ \begin{array}{r} 1 \\ 0  \end{array}\right]\right)= \left[ \begin{array}{r} 2 \\ 0  \end{array}\right] \hspace{1in}  T\left(\left[ \begin{array}{r} 0 \\ 1  \end{array}\right]\right)= \left[ \begin{array}{r} 0 \\ 1  \end{array}\right]  
\end{equation}
$$

The matrix corresponding to the transformation is then

$$
\begin{equation}
A = \left[ \begin{array}{cc} 2 & 0 \\ 0 & 1 \end{array}\right]
\end{equation}
$$

The image of a general vector with components $c_1$ and $c_2$ is given by a linear combination 
of the images of the basis vectors.

$$
\begin{equation}
T\left(\left[ \begin{array}{r} c_1 \\ c_2  \end{array}\right]\right) = T\left(c_1\left[ \begin{array}{r} 1 \\ 0  \end{array}\right] + c_2 \left[ \begin{array}{r} 0 \\ 1  \end{array}\right]\right) = c_1T\left(\left[ \begin{array}{r} 1 \\ 0  \end{array}\right]\right) + c_2 T\left(\left[ \begin{array}{r} 0 \\ 1  \end{array}\right]\right)  = c_1\left[ \begin{array}{r} 2 \\ 0  \end{array}\right] + c_2\left[ \begin{array}{r} 0 \\ 1  \end{array}\right]  = \left[ \begin{array}{cc} 2 & 0 \\ 0 & 1 \end{array}\right] \left[\begin{array}{c} c_1 \\ c_2 \end{array}\right]
\end{equation}
$$

To get a visual understanding of the transformation, we produce a plot that displays multiple input vectors, and their corresponding images.  

**Since we will multiply each vector by the same matrix, we can organize our calculations by constructing a matrix `coords` with columns that match the input vectors.**  

This matrix is the array labeled `coords` in the code below.  We use array slicing to label the first row $\texttt{x}$, and the second row $\texttt{y}$ for the purposes of plotting. 

```python
%matplotlib inline  

import matplotlib.pyplot as plt

# Matrix with columns that match the input vectors
coords = np.array([[0,0],[0.5,0.5],[0.5,1.5],[0,1],[0,0]])
coords = coords.transpose()
coords
x = coords[0,:]
y = coords[1,:]

# Define the Matrix Transformation
A = np.array([[2,0],[0,1]])


# Get a new matrix with columns that match the output vectors
A_coords = A@coords

A_coords
```

The columns of the array A_coords are the images of the vectors that make up the columns of $\texttt{coords}$.  Let's plot the coordinates and see the effect of the transformation.  Again, we need to access the first and second rows by slicing.  The red points show the original coordinates, and the new coordinates are shown in blue.

```python
x_LT1 = A_coords[0,:]
y_LT1 = A_coords[1,:]

# Create the figure and axes objects
fig, ax = plt.subplots()

# Plot the points.  x and y are original vectors, x_LT1 and y_LT1 are images
ax.plot(x,y,'ro')
ax.plot(x_LT1,y_LT1,'bo')

# Connect the points by lines
ax.plot(x,y,'r',ls="--")
ax.plot(x_LT1,y_LT1,'b')

# Edit some settings 
ax.axvline(x=0,color="k",ls=":")
ax.axhline(y=0,color="k",ls=":")
ax.grid(True)
ax.axis([-2,2,-1,2])
ax.set_aspect('equal')
ax.set_title("Horizontal Stretch");
```

We observe that the result of the transformation is that the polygon has been stretched in the horizontal direction.  

### Example 4:  Reflection

Now that we have a set of vectors stored in the $\texttt{coords}$ matrix, let's look at another transformation that is defined by the matrix $B$.

$$
\begin{equation}
B = \left[ \begin{array}{cc} -1 & 0 \\ 0 & 1 \end{array}\right]
\end{equation}
$$

We make a new array $B$, but reuse the plotting code from the previous example.

```python
B = np.array([[-1,0],[0,1]])
B_coords = B@coords

x_LT2 = B_coords[0,:]
y_LT2 = B_coords[1,:]

# Create the figure and axes objects
fig, ax = plt.subplots()

# Plot the points.  x and y are original vectors, x_LT1 and y_LT1 are images
ax.plot(x,y,'ro')
ax.plot(x_LT2,y_LT2,'bo')

# Connect the points by lines
ax.plot(x,y,'r',ls="--")
ax.plot(x_LT2,y_LT2,'b')

# Edit some settings 
ax.axvline(x=0,color="k",ls=":")
ax.axhline(y=0,color="k",ls=":")
ax.grid(True)
ax.axis([-2,2,-1,2])
ax.set_aspect('equal')
ax.set_title("Reflection about y-axis");
```

We see that this transformation reflects points around the vertical axis.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴   Question 3a (2 pts) 

Find a matrix that represents the reflection about the 𝑥-axis.



```python
x_reflect = ...
```

```python
grader.check("3a")
```

<!-- BEGIN QUESTION -->

## 🔴   Question 3b (2 pts) 

Multiply the matrix from part 3a by 𝚌𝚘𝚘𝚛𝚍𝚜 (save the result in the variable `x_reflect_coords`) and plot both the original shape described by 𝚌𝚘𝚘𝚛𝚍𝚜 and the results of the transformation. 


Adjust the axis settings to make sure you can see the original preimage and the new image in your plot.

Label your plot describing this transformatino


```python
x_reflect_coords = ...

...

# Your plotting code above this line
```

```python
grader.check("3b")
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴   Question 4a (2 pts) 

Find a matrix that represents the reflection about the line $y=x$.



```python
x_y_reflect = ...
```

```python
grader.check("4a")
```

<!-- BEGIN QUESTION -->

## 🔴   Question 4b (2 pts) 

Multiply the matrix from part 4a by 𝚌𝚘𝚘𝚛𝚍𝚜 (save the result in the variable `x_y_reflect_coords`) and plot both the original shape described by 𝚌𝚘𝚘𝚛𝚍𝚜 and the results of the transformation. 

Adjust the axis settings to make sure you can see the original preimage and the new image in your plot.

Label your plot describing this transformation.


```python
x_y_reflect_coords = ...

...

# Your plotting code above this line
```

```python
grader.check("4b")
```

<!-- END QUESTION -->



### Example 5:  Rotation

To rotate vectors in the plane, we choose an angle $\theta$ (in radians) and write down the matrix that represents the rotation counterclockwise by an angle $\theta$.  Basic trigonometry can be used to calculate the columns in this case.   

<img src="img/rotation.png" width = 800> 

Thus, the standard matrix for a counterclockwise rotation is given by:   

$$
\begin{equation}
R = \left[ \begin{array}{cc} \cos\theta & -\sin\theta \\ \sin\theta & \cos\theta \end{array}\right]
\end{equation}
$$

Note:  All angles should be input in radians

**Do NOT use python `math` functions** (they will cause the Gradescope grader to fail).  

Instead use Numpy functions:
`np.pi`, `np.cos()`, `np.sin()`, `np.sqrt()` etc

```python
# Heres's an example of a counterclockwise rotation of pi/6:

theta = np.pi/6
R = np.array([[np.cos(theta),-np.sin(theta)],[np.sin(theta),np.cos(theta)]])
R_coords = R@coords

x_LT3 = R_coords[0,:]
y_LT3 = R_coords[1,:]

# Create the figure and axes objects
fig, ax = plt.subplots()

# Plot the points.  x and y are original vectors, x_LT1 and y_LT1 are images
ax.plot(x,y,'ro')
ax.plot(x_LT3,y_LT3,'bo')

# Connect the points by lines
ax.plot(x,y,'r',ls="--")
ax.plot(x_LT3,y_LT3,'b')

# Edit some settings 
ax.axvline(x=0,color="k",ls=":")
ax.axhline(y=0,color="k",ls=":")
ax.grid(True)
ax.axis([-2,2,-1,2])
ax.set_aspect('equal')
ax.set_title("Rotation");
```

### Example 6: Shear

In the study of mechanics, shear forces occur when one force acts on part of a body while another force acts on a different part of the body in the opposite direction. To visualize this, imagine a deck of playing cards resting on a table. If you place your hand on top of the deck and slide it parallel to the table, the cards shift relative to one another—this is a shear effect.

A shear transformation is a linear transformation that shifts points in a fixed direction by an amount proportional to their (signed) distance from a line that remains fixed. This line is parallel to the direction of motion and passes through the origin.

For example, a horizontal shear leaves the y-coordinate unchanged and adds to the x-coordinate an amount proportional to the y-coordinate (subtracting when y is negative).

Horizontal Shear:
$$\left[ \begin{array}{r} 1 & k \\ 0 & 1  \end{array}\right]  $$

where  $k \in \mathbb{R}$ is called the **shearing factor**.

Similarly, here is the standard matrix for a vertical Shear: $$
\left[ \begin{array}{r} 1 & 0 \\ k & 1  \end{array}\right]$$

```python
#horizontal shear
S = np.array([[1,2],[0,1]])
S_coords = S@coords

x_LT4 = S_coords[0,:]
y_LT4 = S_coords[1,:]


# Create the figure and axes objects
fig, ax = plt.subplots()

# Plot the points.  x and y are original vectors, x_LT1 and y_LT1 are images
ax.plot(x,y,'ro')
ax.plot(x_LT4,y_LT4,'bo')

# Connect the points by lines
ax.plot(x,y,'r',ls="--")
ax.plot(x_LT4,y_LT4,'b')

# Edit some settings 
ax.axvline(x=0,color="k",ls=":")
ax.axhline(y=0,color="k",ls=":")
ax.grid(True)
ax.axis([-2,4,-1,2])
ax.set_aspect('equal')
ax.set_title("Shear");
```

### Example 7:  Composition of transformations

One of the powerful aspects of representing linear transformations as matrices is that we can form compositions by matrix multiplication.  For example,  in order to find the matrix that represents the composition $B\circ R$, the rotation from Example 5 (30 degrees), followed by the reflection of Example 4 (about the y-axis),
we can simply multiply the matrices that represent the individual transformations.  

$$
\begin{equation}
BR = \left[ \begin{array}{cc} -1 & 0 \\ 0 & 1 \end{array}\right]
\left[ \begin{array}{cc} \cos\theta & -\sin\theta \\ \sin\theta & \cos\theta \end{array}\right]
=\left[ \begin{array}{cc} -\cos\theta & \sin\theta \\ \sin\theta & \cos\theta \end{array}\right]
\end{equation}
$$


```python
theta = np.pi/6

C = np.array([[-1,0],[0,1]])@np.array([[np.cos(theta),-np.sin(theta)],[np.sin(theta),np.cos(theta)]])
C_coords = C@coords

x_LT5 = C_coords[0,:]
y_LT5 = C_coords[1,:]


# Create the figure and axes objects
fig, ax = plt.subplots()

# Plot the points.  x and y are original vectors, x_LT1 and y_LT1 are images
ax.plot(x,y,'ro')
ax.plot(x_LT5,y_LT5,'bo')

# Connect the points by lines
ax.plot(x,y,'r',ls="--")
ax.plot(x_LT5,y_LT5,'b')

# Edit some settings 
ax.axvline(x=0,color="k",ls=":")
ax.axhline(y=0,color="k",ls=":")
ax.grid(True)
ax.axis([-2,2,-1,2])
ax.set_aspect('equal')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴   Question 5a (2 pts) 

Find a matrix that represents a vertical shear with shearing factor 3 followed by a 60 degree counterclockwise rotation.

(Do NOT use python `math` functions (they will cause the Gradescope grader to fail).  
Instead use Numpy functions:
`np.pi`, `np.cos()`, `np.sin()` etc)

```python

shear_rot = ...


```

```python
grader.check("5a")
```

<!-- BEGIN QUESTION -->



## 🔴   Question 5b (2 pts) 

Find a matrix that represents a 60 degree counterclockwise rotation followed by a vertical shear with shearing factor 3 

```python

rot_shear = ...
```

```python
grader.check("5b")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## 🔴   Question 5c (3 pts) 

Apply the transformation from 5a to `coords` and assign the resulting matrix to the variable `shear_rot_coords`.

Then apply the transformation from 5b to `coords` and assign the resulting matrix to the variable `rot_shear_coords`.

Then plot both of these transformations along with the image represented by coords. 
Use blue to represent the transformation from 5a and green to represent the transformation from 5b.

Adjust the axis settings to make sure you can see the original preimage and the new image in your plot.

Label your plot describing these transformations.


```python
shear_rot_coords = ...
rot_shear_coords = ...

print(shear_rot_coords)

...

# Your plotting code above this line
```

```python
grader.check("5c")
```

<!-- END QUESTION -->



## 🔴   Question 5d (2 pts) 

Construct a matrix that represents a horizontal and a vertical stretch by a factor of 2 .



```python
hor_ver_stretch = ...
```

```python
grader.check("5d")
```

<!-- BEGIN QUESTION -->

## 🔴   Question 5e (2 pts) 

Apply the transformation from Part 5d to `coords`.  Save the result in the variable `hor_ver_stretch_coords` and plot. 

Adjust the axis settings to make sure you can see the original preimage and the new image in your plot.

Label your plot describing this transformation.


```python
hor_ver_stretch_coords = ...

...

# Your plotting code above this line
```

```python
grader.check("5e")
```

<!-- END QUESTION -->

## Applications of Linear Transformations

### Computer graphics

Computer graphics essentially consist of a large number of basic geometric shapes such as triangles or quadrilaterals, defined by their vertexes (corners).  The ability to manipulate these shapes efficiently is a fundamental necessity in the field of computer graphics.  Linear transformations can be used to generate complex shapes from basic shapes through effects such as shearing and reflection, and can also be used to rotate and translate shapes in order to create a sequence of images needed for an animated scene.

The software responsible for rendering three-dimensional scenes in modern video games must process the coordinates of millions of vertexes in each frame of animation.  Typical hardware used for running the software will be able to generate around 60 frames per second.  In addition to the placement of vertexes in a scene, there are other computations required, such as those to provide color and shading, and then to project the three-dimensional scene onto the two-dimensional screen.  In this context, the efficiency of the algorithms, software, and hardware all play a major role in the quality of the end result.  

The use of matrix multiplication to manipulate coordinates for computer graphics has two primary advantages.  Matrix representations can be used to merge multiple transformations together into a single transformation.  For example, instead of applying a rotation followed by a scaling transformation, the two operations can easily be merged into a single operation by multiplying the associated matrices as we did earlier.  The other advantage to organizing the calculations into matrix multiplications is that dedicated hardware (graphics processing units) is designed specifically to handle these calculations in large numbers.

The construction of matrices to carry out rotations, reflections, and shears, has already been demonstrated in the exercises above.  Here we show how to use matrix multiplication to carry out a translation, which is not a linear transformation.  

#### Translation

A **translation** is a transformation that moves all points an equal amount in the same direction.  Shown below is an example where all points are shifted (translated) three units to the right, and one unit down by a transformation $T:\mathbb{R}^2 \to \mathbb{R}^2$.  In the plot, we show several points which define a shape, and their corresponding images under the translation.  

```python
%matplotlib inline  

# Define a matrix of coordinates that trace out an L shape:
coords_L = np.array([[0,0],[0,3],[1,3],[1,1],[2,1],[2,0],[0,0]])
# Let each column represent a vector that starts at the origin and ends at a point on the shape.
coords_L = coords_L.transpose()

x = coords_L[0,:]
y = coords_L[1,:]

## Compute translation by adding 3 to x coordinates and -1 to y coordinates
x_translated =  np.copy(x)
y_translated =  np.copy(y)

for i in range(x.shape[0]):
    x_translated[i] += 3
    y_translated[i] += -1
```

```python
# Create the figure and axes objects
fig, ax = plt.subplots()

# Plot the points.
ax.plot(x,y,'ro')
ax.plot(x_translated,y_translated,'bo')

# Connect the points by lines
ax.plot(x,y,'r',ls="--")
ax.plot(x_translated,y_translated,'b')

# Edit some settings 
ax.axvline(x=0,color="k",ls=":")
ax.axhline(y=0,color="k",ls=":")
ax.grid(True)
ax.axis([-1,6,-2,6])
ax.set_aspect('equal')
ax.set_title("Translation");
```

Translation of coordinates is useful in computer graphics, but it is not a linear transformation.  

In order to see why, we write a formula for the image of a vector $\mathbf{x}$, and then verify that $T(\mathbf{x+y}) \neq T(\mathbf{x})+T(\mathbf{y})$ for arbitrary vectors $\mathbf{x}$ and $\mathbf{y}$.

$$
\begin{equation}
T(\mathbf{x}) = \mathbf{x} + \left[ \begin{array}{r} 3 \\ -1 \end{array} \right]
\end{equation}
$$


$$
\begin{equation}
T(\mathbf{x}+\mathbf{y}) = (\mathbf{x} + \mathbf{y}) + \left[ \begin{array}{r} 3 \\ -1 \end{array} \right]
\end{equation}
$$

$$
\begin{equation}
T(\mathbf{x}) + T(\mathbf{y}) = \mathbf{x} + \left[ \begin{array}{r} 3 \\ -1 \end{array} \right] + \mathbf{y} + \left[ \begin{array}{r} 3 \\ -1 \end{array} \right]
= (\mathbf{x} + \mathbf{y}) + \left[ \begin{array}{r} 6 \\ -2 \end{array} \right]
\end{equation}
$$

$$\implies T(\mathbf{x}+\mathbf{y})  \neq T(\mathbf{x}) + T(\mathbf{y})$$

The fact that $T$ is **not linear** implies that this transformation **cannot be represented as a matrix multiplication in its current state.**

Since the use of translations are necessary, a standard method has been developed to incorporate matrix multiplication and the associated advantages.  The method requires that we view $\mathbb{R}^2$ as a plane within $\mathbb{R}^3$.  We can then construct a linear transformation $L:\mathbb{R}^3\to\mathbb{R}^3$ which will have the *effect of the translation* within that plane.  We demonstrate the details for the transformation defined above.  

First, we associate the coordinates $(x_1,x_2)$ of a point in $\mathbb{R}^2$ with the coordinates $(x_1,x_2,1)$ of a point in $\mathbb{R}^3$ in the plane $x_3=1$.  These new coordinates are known as **homogeneous coordinates**.  We can then create a linear transformation $L:\mathbb{R}^3\to\mathbb{R}^3$ that represents a shear that is parallel to the $x_1x_2$-plane, and in the direction of the desired translation.  The effect of $L$ on points in the plane $x_3=1$ is a translation within the $x_3=1$ plane.  
\begin{equation}
A = \left[ \begin{array}{rrr} 1 & 0 & \text{Translation in x direction} \\ 0 & 1 & \text{Translation in y direction} \\ 0 & 0 & 1 \end{array} \right]
\end{equation}



For example, for the translation showed in the picture above, we define $L(X)=AX$, where $A$ is the following matrix.


$$
\begin{equation}
A = \left[ \begin{array}{rrr} 1 & 0 & 3 \\ 0 & 1 & -1 \\ 0 & 0 & 1 \end{array} \right]
\end{equation}
$$

Recall that the columns of $A$ are the images of the standard basis vectors $\{e_1, e_2, e_3\}$.  In this case $e_1$ and $e_2$ are unchanged, and $e_3$ is sheared in the desired direction.  We observe that $L$ has the intended effect.


$$
\begin{equation}
L(X) = AX = \left[ \begin{array}{rrr} 1 & 0 & 3 \\ 0 & 1 & -1 \\ 0 & 0 & 1 \end{array} \right]
\left[ \begin{array}{r} x_1 \\ x_2 \\ 1 \end{array} \right] = 
\left[ \begin{array}{c} x_1+3 \\ x_2-1 \\ 1 \end{array} \right]
\end{equation}
$$

Finally, we associate with $L(X)$ the point in $\mathbb{R}^2$ that has coordinates $(x_1+3, x_2-1)$.

Let's have a look at the revised calculation with NumPy, that now makes use of the matrix multiplication.

```python
# Form a matrix of coordinates with rows x, y, 1

coords_L = np.array([[0,0],[0,3],[1,3],[1,1],[2,1],[2,0],[0,0]])
coords_L = coords_L.transpose()

x = coords_L[0,:]
y = coords_L[1,:]

Ones = np.ones((1,7))
new_coords = np.vstack((x,y,Ones))

# Define the matrix for the transfomation L
A = np.array([[1,0,3],[0,1,-1],[0,0,1]])

## Compute translation by matrix multiplication
coords_translated = A@new_coords

## Slice off the first and second rows.  These are the coords of the translated points
x_translated = coords_translated[0,:]
y_translated = coords_translated[1,:]

# Create the figure and axes objects
fig, ax = plt.subplots()

# Plot the points.
ax.plot(x,y,'ro')
ax.plot(x_translated,y_translated,'bo')

# Connect the points by lines
ax.plot(x,y,'r',ls="--")
ax.plot(x_translated,y_translated,'b')

# Edit some settings 
ax.axvline(x=0,color="k",ls=":")
ax.axhline(y=0,color="k",ls=":")
ax.grid(True)
ax.axis([-1,6,-1,6])
ax.set_aspect('equal')
ax.set_title("Translation");
```

If we look closely at the computation, we realize that this way of computing the coordinates of the translated shape actually requires *more operations* than the first attempt where we stepped through the array and manipulated the coordinates directly.  The important point to remember however is that the **ability to represent the translation as a matrix multiplication means that we can merge it with another operation, such as rotation, in order to carry out both operations by multiplying by a single matrix**.

Translations of shapes in three dimensions can be done in a similar way by associating each set of coordinate $(x_1, x_2, x_3)$ in $\mathbb{R}^3$ with the point $(x_1, x_2, x_3, 1)$ in $\mathbb{R}^4$.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴   Question 6a (2 pts) 

Find a single matrix that represents a transformation that has the effect of a reflection about the line $y=x$
 followed by a shift four units to the left.  Assign this to the variable `x_y_reflect_shift`



```python
x_y_reflect_shift = ...

```

```python
grader.check("6a")
```

<!-- BEGIN QUESTION -->

## 🔴   Question 6b (2 pts) 

Apply the transformation from part (a) to a corresponding matrix representing the L shape defined by `coords_L` and plot the resulting image in the xy-plane. (Use homogenous coordinates and matrix multiplication).

Adjust the axis settings to make sure you can see the original preimage and the new image in your plot.

Label your plot describing this transformation.


```python
...
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴   Question 7a (2 pts) 

Find a single matrix that represents a $45$ degree counterclockwise rotation about the point $(1,2)$.
Assign your answer to the variable `question7a`


Hint: (In order to carry out the rotation about the point $(1,2)$, we first need to bring the center of rotation to (0,0), for which we need to do translation involving the shift of 1
 unit left and 2
 units down.  Then we need to perform the desired rotation, and then translate back up one unit to the right and two units up). 

Again, use Numpy functions for any cos, sin or square roots (do not use Python's Math library)




```python


...

question7a = ...


```

```python
grader.check("7a")
```

<!-- BEGIN QUESTION -->

## 🔴   Question 7b (2 pts) 

Apply the transformation from part (a) to a corresponding matrix representing the L shape defined by `coords_L` and plot the resulting image in the xy-plane. (Use homogenous coordinates and matrix multiplication).

Adjust the axis settings to make sure you can see the original preimage and the new image in your plot.

Label your plot describing this transformation.


```python


...

# Your plotting code above this line
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴   Question 8a (3 pts) 

Find a matrix that represents **clockwise** rotation of $5\pi/6$ radians about the point $(1,1)$ followed by a shift $3$ units to the right.  Assign this matrix to the variable `question8a`




```python
...

question8a = ...



```

```python
grader.check("8a")
```

<!-- BEGIN QUESTION -->

## 🔴   Question 8b (2 pts) 

Apply the transformation from part (a) to a corresponding matrix representing the L shape defined by `coords_L` and plot the resulting image in the xy-plane. (Use homogenous coordinates and matrix multiplication).


Label your plot describing this transformation.



```python


...

# Your plotting code above this line
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴   Question 9a (2 pts) 

Find a single matrix that represents a transformation that has the effect of reflection about the $y$-axis followed by a shift $2$
 units to the right and $3$
 units up.  Assign this to the variable `question9a`

```python


question9a = ...

```

```python
grader.check("9a")
```

<!-- BEGIN QUESTION -->

## 🔴  Question 9b (2 pts) 

Apply the transformation from part (a) to a corresponding matrix representing the L shape defined by `coords_L` and plot the resulting image in the xy-plane. (Use homogenous coordinates and matrix multiplication).

Adjust the axis settings to make sure you can see the original preimage and the new image in your plot.

Label your plot describing this transformation.

```python


...

# Your plotting code above this line
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


### Submission Instructions


Before proceeding any further, **save this notebook.**

## IMPORTANT STEP BEFORE COMPILING THIS ASSIGNMENT
To **make sure all of your plots render correctly in the PDF,** please do the following steps:

1.) In the toolbar, go to Settings  
2.) Disable Autosave Documents 

3.) From the toolbar select Kernel-> Restart Kernel and Run Up to Selected Cell (up to this cell) 

4.) Check that all plots are correctly displayed in this notebook and then Save.  

AFTER running the last cell below, open your PDF and **check that your pdf has all plots correctly showing before uploading to Gradescope - otherwise you WILL lose points**

**TROUBLESHOOTING TIPS**  If you are having any issues compiling your assignment, please [read through these troubleshooting tips first](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing), then post any questions on Piazza.  




Then run the cell below to double check that you don't have any spaces between dollar signs and text when writing LaTeX:

```python
# Run this cell before you run the 'grader.export()' cell below.  
# It will search for LaTeX errors that will cause the LaTeX compiler to fail.  

import simple_latex_checker as slc

nb = slc.Nb_checker()
nb.run_check("lab02.ipynb")
```

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 


* **Lab 2 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the Lab 2 Autograded assignment in Gradescope.

* **Lab 2 Manually Graded**: Submit your lab02.PDF to the Lab 2 Manually Graded assignment in Gradescope.  **It is your responsibility to fully review your PDF file before submitting and make sure that all your lines of code are visible and any LaTeX has correctly compiled and is fully viewable.**  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.** If not, you will lose points.    

[TROUBLESHOOTING TIPS](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing)  If you are having any issues compiling your assignment, please read through these troubleshooting tips first, then post any questions on Piazza.  

**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='lab02.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for 2a (points: 3)
```
>>> assert get_array_hash_normalized(S_stand_matrix) == '81a8eb7d26d577393b3523e7d396cc76'
```
```
>>> assert np.array_equal(R_composed_with_S(np.array([[1], [4], [-2]])), np.array([[15], [-11]]))
```
```
>>> assert np.array_equal(R_composed_with_S(np.array([[0], [0], [0]])), np.array([[0], [0]]))
```


Tests for 3a (points: 2)
```
>>> assert get_array_hash_normalized(x_reflect) == 'bdad812cb893b39f7f80369512305d31'
```


Tests for 3b (points: 1)
```
>>> assert get_array_hash_normalized(x_reflect_coords) == '3b1b8d57618b146592ba6e8bf228609f'
```


Tests for 4a (points: 2)
```
>>> assert get_array_hash_normalized(x_y_reflect) == '3aa0cfa677fff445815da99720737752'
```


Tests for 4b (points: 1)
```
>>> assert get_array_hash_normalized(x_y_reflect_coords) == '979bce9fe12350176d97ecf44ecdfc17'
```


Tests for 5a (points: 2)
```
>>> assert get_array_hash_normalized(shear_rot.round(7)) == 'c00895b8c8cb255c57593962e2cc730f'
```


Tests for 5b (points: 1)
```
>>> assert get_array_hash_normalized(rot_shear.round(7)) == '3ee61ca5a5bc25134d01a617f1b758e1'
```


Tests for 5c (points: 1)
```
>>> assert get_array_hash_normalized(shear_rot_coords.round(7)) == '2ff3defd43ff6402fd7138ba2eff6f4e'
```
```
>>> assert get_array_hash_normalized(rot_shear_coords.round(7)) == '4349c7d1b503decea48a3e872ff55fbe'
```


Tests for 5d (points: 2)
```
>>> assert get_array_hash_normalized(hor_ver_stretch) == '17d1a2c8408926e2c9dd0f886d9601d7'
```


Tests for 5e (points: 1)
```
>>> assert get_array_hash_normalized(hor_ver_stretch_coords) == 'fa2914b9ca202db08b3b60b25a6ae800'
```


Tests for 6a (points: 2)
```
>>> assert get_array_hash_normalized(x_y_reflect_shift) == '65f921876250240b27b69f5b77a5094a'
```


Tests for 7a (points: 2)
```
>>> assert get_array_hash_normalized(question7a.round(7)) == 'f0a1c5536a56c65e845dffd997cb26c6'
```


Tests for 8a (points: 3)
```
>>> assert get_array_hash_normalized(question8a.round(7)) == '6357a98d92d46ab867b58a08deef19dc'
```


Tests for 9a (points: 2)
```
>>> assert get_array_hash_normalized(question9a) == '6e74187d46a95ad5a4d7c574d560b06b'
```


Tests for q1a (points: 1)
```
>>> assert get_array_hash_normalized(T_of_V) == '49738dd130e35d76ffee1a69c810763f'
```


Tests for q1b (points: 1)
```
>>> assert get_array_hash_normalized(U) == 'b7de76572bad455a78c21c6ea10fb153'
```


Tests for q1c (points: 2)
```
>>> assert np.array_equal(N(np.array([[5], [-1], [2]])), np.array([[-8], [10]]))
```
```
>>> assert np.array_equal(N(np.array([[0], [0], [0]])), np.array([[0], [3]]))
```


Tests for q1e (points: 2)
```
>>> assert get_array_hash_normalized(T_stand_matrix) == '4361fd79cbc4553be6f815026e2e789b'
```


Tests for q1f (points: 2)
```
>>> assert get_array_hash_normalized(L_T_stand_matrix) == 'ba1692ae467c73affbaf76f101405d76'
```


Tests for q2b (points: 1)
```
>>> assert get_hash(ans_2b) == '9d5ed678fe57bcca610140957afab571'
```
