```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

<a id='verytop'></a>

# Lab 1: Applications of Systems of Linear Equations

## Due Date: Thurs, Jan 22nd 11:59 PM on Gradescope


### Detailed Submission Instructions Are Provided at the end of this Notebook


## Collaboration Policy

A key step in learning and retention is **creating solutions on your own.**   Below are examples of acceptable vs unacceptable use of resources and collaboration when doing lab assignments in CSCI 2820.


The following would be some **examples of cheating** when working on HW assignments in CSCI 2820.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given lab problem or one similar is cheating.  However, consulting notes from the class videos, and web pages that explain the material taught in class but do NOT show a solution to the lab problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course videos, the course textbooks, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the lab before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a lab problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a lab solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit office Hours and/or post questions on Piazza!



## Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


### SymPy

**SymPy** is an open-source Python library designed for symbolic mathematics. It functions as a full-featured computer algebra system (CAS) within the Python environment, allowing users to perform various mathematical operations symbolically rather than numerically. 

```python
# import SymPy with alias sp
import sympy as sp


# Function needed to run in-notebook tests
import hashlib

from sympy.printing.repr import srepr

def get_hash(expr):
    """Hash a SymPy expression canonically."""
    expr = sp.nsimplify(expr)          # 3.0 → 3, 0.5 → 1/2
    return hashlib.md5(sp.srepr(expr).encode()).hexdigest()

def hash_simplified_matrix(M):
    M = sp.Matrix(M).applyfunc(sp.nsimplify)
    return hashlib.sha256(sp.srepr(M).encode("utf-8")).hexdigest()

def get_hash_symb(expr):
    expr = sp.sympify(expr)

    # If there are floats anywhere, try to convert them to exact rationals
    expr = expr.replace(
        lambda e: isinstance(e, sp.Float),
        lambda e: sp.nsimplify(e)
    )

    expr = sp.simplify(expr)
    return hashlib.md5(srepr(expr).encode("utf-8")).hexdigest()



def hash_simplified_matrix_floats(M):
    M = sp.Matrix(M)

    def canon(e):
        e = sp.sympify(e)
        # convert floats like 0.3333333333 to Rational(1, 3) when possible
        if e.is_Float:
            e = sp.nsimplify(e)
        return sp.simplify(e)

    M = M.applyfunc(canon)
    return hashlib.sha256(srepr(M).encode("utf-8")).hexdigest()
```

We can represent the matrix 


   $A=\begin{bmatrix}
   1  & 2  & 3  \\
   4  & 5 &  6 \\
  \end{bmatrix}$




in SymPy as follows:


```python
# Run this cell (shift-enter)
A = sp.Matrix([[1,2,3],[4,5,6]])
A
```

### Gauss-Jordan Elimination

The complete process of applying row operations to reduce an augmented matrix to 
reduced row echelon form can be expressed as a recursive process in an algorithmic fashion, making
it possible to program computers to solve linear systems. 

This algorithm is called Gauss-Jordan elimination.


Here are the steps to do so:

 - **Step 1:** Begin with the leftmost nonzero column (if there is one). This will be a pivot column.

 - **Step 2:** Select a nonzero entry in this pivot column as a pivot. If necessary, interchange rows to
move this entry to the first row (this entry will be a pivot).

 - **Step 3:** Use row operations to create zeros in all positions below the pivot.

 - **Step 4:** Cover (or ignore) the row containing the pivot position and cover all rows, if any, above
it. Apply steps 1-3 to the submatrix that remains. Repeat the process until there are no more
nonzero rows to modify.
To obtain the reduced row echelon form we need one more step.

 - **Step 5:** Beginning with the rightmost pivot and working upward and to the left, create zeros above
each pivot. If a pivot is not 1, make it 1 by an appropriate row multiplication.
The algorithm described in steps 1-4 will produce the row echelon form of the matrix. This
algorithm is called Gaussian elimination. 

When we add step 5 to produce the reduced row echelon
form, the algorithm is called **Gauss-Jordan elimination**.



SymPy has a built-in method that uses Gauss-Jordan elimination to find the RREF of a matrix.  
Click on this link to read the documentation of the `.rref()` method in SymPy
(https://docs.sympy.org/latest/modules/matrices/matrices.html#sympy.matrices.matrixbase.MatrixBase.rref)
                                  

```python
#Run this cell to find the RREF of matrix A from above.
A.rref()
```

The first entry in the list returned by the `.rref()` method is the RREF matrix:

```python
# Run this cell to output the RREF matrix
A.rref()[0]
```

The 2nd entry returned is the indices of columns that contain a pivot (starting at 0): 

```python
# Run this cell to output the incides of the columns that contain a pivot
A.rref()[1]
```

Alternatively, you can just return the RREF by using the argument pivots = False:

```python
# Run this cell to see an alternative way to output the RREF matrix without the list of pivots
A.rref(pivots = False)
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴 Question 1 ab (2 points)   



Consider the system of linear equations:
\begin{align*} 
x+2y-z &=  1 \\ 
3x + 2y+2z &=  7 \\
-x+ 4z &= -3
\end{align*}

a). Write this system as an augmented matrix using SymPy. Assign it to the variable `q1_aug` (replace the ellipses ... with your solution)

b). Use SymPy to find the RREF of your augmented matrix.  Save the RREF matrix to the variable named `q1_rref`. 

```python
#Replace ... with your answer

q1_aug = ...

# This next line will print out your matrix after you've assigned it
q1_aug
```

```python
#Replace ... with your answer

q1_rref = ...

q1_rref

```

```python
grader.check("q1ab")
```

## 🔴 Question 1c (1 pt)  

**Multiple Choice:**
Based on the RREF matrix you found above, choose the answer that best describes the system of linear equations and then assign the the variable  `ans_1c` to the corresponding letter (as a capital letter string `'A'`,`'B'`, or `'C'`) to  below to indicate your answer.  Replace the ellipses (`...`) with your answer. 

A).  The system is inconsistent

B).  The system is consistent and the solution is unique

C).  The system is consistent and the solution is not unique.  

```python
ans_1c = ...


ans_1c

```

```python
grader.check("q1c")
```

## 🔴 Question 1d 

Use the RREF to describe the solution set to the original system of linear equations. 

Use SymPy's solution dictionary to write your answer.  That is, your code should have the form:



```python

x, y, z  = sp.symbols('x,y,z')
sol = {
    x: answer,
    y: answer,
    z: answer
}

```

If no solution exists, leave the dictionary empty, that is sol = { } 

```python
# Enter the solution to the system of equations here:

x, y, z  = sp.symbols('x,y,z')

# Write your answer as a dictionary, using the format shown above:

sol1 = ...


```

```python
grader.check("q1d")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴 Question 2 (4 points)   



Consider the system of linear equations:
\begin{aligned}
-2x_1 + 4x_2 + 2x_3 - 8x_4 + 4x_5 &= -8 \\
 3x_1 - 6x_2 - 2x_3 + 11x_4- 7x_5 &= 13 \\
  x_1 - 2x_2 - 5x_3 + 8x_4 +  x_5 &= -3
\end{aligned}

a). Write this system as an augmented matrix using SymPy. Assign it to the variable `q2_aug`. 

b). Use SymPy to find the RREF of your augmented matrix.  Save the RREF matrix to the variable named `q2_rref`. 

c).  Based on the RREF matrix you found above, choose the answer that best describes the system of linear equations and then assign the the variable  `ans_2c` to the corresponding letter (as a capital letter string `'A'`,`'B'`, or `'C'`) to  below to indicate your answer.  Replace the ellipses (`...`) with your answer. 

     A).  The system is inconsistent

     B).  The system is consistent and the solution is unique

     C).  The system is consistent and the solution is not unique.  

d). Write the solution set to the original system of linear equations.  If you have free variables don't rename them (i.e. use x1, x2, etc).   Use SymPy's solution dictionary to write your answer.  That is, your code should have the form:



```python

x1, x2, x3, x4, x5  = sp.symbols('x1 x2 x3 x4 x5')
sol = {
    x1: answer(s),
    x2: answer(s),
    x3: answer(s),
    x4: answer(s), 
    x5: answer(s)
}


If no solution exists, leave the dictionary empty, that is sol = { } 

```python
# Replace the '...' with your solution to part (a)

q2_aug = ...

q2_aug
```

```python
# Replace the '...' with your answer to part (b)


q2_rref = ...

q2_rref
```

```python
# Replace the '...' with your answer to part (c)


ans_2c = ...
```

```python
# Define a SymPy dictionary giving the general solution.  Do not rename any free variables. 

x1,x2,x3,x4,x5 = sp.symbols('x1 x2 x3 x4 x5')

sol2 = ...

```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴   Question 3  


A mining company has three mines. 

One day of operation at the mines produces the following output:

    
 - Mine 1 produces 25 tons of copper, 600 kilograms of silver and 15 tons of manganese.
    
 - Mine 2 produces 30 tons of copper, 500 kilograms of silver and 10 tons of manganese.
    
 - Mine 3 produces 20 tons of copper, 550 kilograms of silver and 12 tons of manganese.
    
Suppose the company has orders for 550 tons of copper, 11350 kilograms of silver and 250 tons of
manganese.



#### Question 3a (2 pts) ###

i).  Write a system of equations to answer the question:   **How many days should the company operate each mine to exactly fill the orders?**

ii).  State clearly what the variables in your system represent.

- Tip:  To write math equations in Markdown use LaTeX.  Double click on the cell in Question 1 with the system of equations to see an example of the code to write a system of equations (you can copy and paste the LaTeX code and modify it below).  



#### Question 3a Solution) 

**Double click on this cell and then type your work answering Question 3ai and 3aii** using LaTeX in this cell.   Do not add any additional cells to this part.  

<!-- END QUESTION -->

#### Question 3 cont'd (3 pts) 

 3b).  Use SymPy to write the system as an augmented matrix 
 
 3c). Then find the RREF 
 
 3d). Then use RREF to find the solution set to the original system of equations.   Let $m_i$ be the number of days the company should operate mine $i$
and write the solution set in the form 
```python

m1, m2, m3  = sp.symbols('m1 m2 m3')
sol3 = {
    m1: answer(s), 
    m2: answer(s),
    m3: answer(s)
}

```
If no solution exists, leave the dictionary empty, that is sol = { } 

```python
q3_aug = ...

q3_rref = ...


```

```python
# Use the RREF to find the solution set to the system of linear equations.  Write your answer as a Sympy dictionary

# Let $m_i$ = the number of days the company should operate mine $i$

m1, m2, m3  = sp.symbols('m1 m2 m3')

sol3 = ...

```

```python
grader.check("q3b")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴  Question 4 (5 pts)  ###


Shown below is the traffic pattern in the downtown area of a large city. The figure
gives the number of cars per hour traveling along each road. Any car that drives
into an intersection must also leave the intersection. This means that the number of
cars entering an intersection in an hour is equal to the number of cars leaving the
intersection.

<img src="img/traffic.png"> 


#### 4a).  Use the picture to write a system of linear equations.


#### Question 4a Solution) 

**Type your work answering to Question 4a** in this cell (use LaTeX).  Show all of your steps and fully justify your answer.  Do not add any additional cells to this part.  

<!-- END QUESTION -->

#### Question 4 cont'd (4 pts) ###

 4b). Write this system as an augmented matrix using SymPy.  **Use the following column orderering for your variables:  $w, x, y, z$**   

 4c). Then find the RREF.

 4d).   Based on the RREF matrix you found above, choose the answer that best describes the system of linear equations and then assign the the variable  `ans_4d` to the corresponding letter (as a capital letter string `'A'`,`'B'`, or `'C'`) to  below to indicate your answer.  Replace the ellipses (`...`) with your answer. 

     A).  The system is inconsistent

     B).  The system is consistent and the solution is unique

     C).  The system is consistent and the solution is not unique.  
                                                                               
                                                                                 
 4e). Use the RREF to find the solution set to the original system of equations.   Do not rename any free variables.
```python

w, x, y, z  = sp.symbols('w x y z')
sol4 = {
    w: answer(s), 
    x: answer(s),
    y: answer(s)
    z: answer(s)
}

```

If no solution exists, leave the dictionary empty, that is sol = { } 

```python
q4_aug = ...


q4_rref = ...

q4_rref
```

```python
# Answer the multiple choice question here
ans_4d = ...

```

```python
# Use the RREF to find the solution set to the system of linear equations.  Write your answer as a Sympy dictionary

w, x, y, z  = sp.symbols('w x y z')

sol4 = ...
```

```python
grader.check("q4b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴  Question 5  ###


In a previous math classes, you have probably seen the fact that, if we are given two points in the plane, then there is a unique line passing through both of them. In this problem, we will begin with the four points on the left below and ask to find a polynomial that passes through these four points as shown on the right.

<img src="img/parabola.png" width="600">

<!-- BEGIN QUESTION -->

A degree three polynomial can be written as
$$p(x) = a+bx+cx^2+dx^3$$
where $a$, $b$, $c$ and $d$ are coefficients that we would like to determine.  


#### 5a). Write the four linear equations for the coefficients obtained by requiring that the graph of the polynomial passes through the four points above.


#### Question 5a Solution) 

**Type your work answering to Question 5a** in this cell (use LaTeX).  Show all of your steps and fully justify your answer.  Do not add any additional cells to this part.  

<!-- END QUESTION -->

#### Question 5 cont'd (4 pts) ###

 5b). Write this system as an augmented matrix using SymPy.  **Use the following column orderering for your variables:  $w, x, y, z$**   

 5c). Then find the RREF.

 5d).   Based on the RREF matrix you found above, choose the answer that best describes the system of linear equations and then assign the the variable  `ans_5d` to the corresponding letter (as a capital letter string `'A'`,`'B'`, or `'C'`) to  below to indicate your answer.  Replace the ellipses (`...`) with your answer. 

     A).  The system is inconsistent

     B).  The system is consistent and the solution is unique

     C).  The system is consistent and the solution is not unique.  
                                                                               
                                                                                 
 5e). Use the RREF to find the solution set to the original system of equations.   Do not rename any free variables.
```python

a, b, c, d  = sp.symbols('a b c d')
sol5 = {
    a: answer(s), 
    b: answer(s),
    c: answer(s)
    d: answer(s)
}

```

If no solution exists, leave the dictionary empty, that is sol = { } 


Check your work by graphing the polynomial(s) you found in Desmos and verifying that it passes through the four points.   https://www.desmos.com/calculator

```python
q5_aug = ...


q5_rref = ...

q5_rref
```

```python
# Answer the multiple choice question here
ans_5d = ...

```

```python
# Use the RREF to find the solution set to the system of linear equations.  Write your answer as a Sympy dictionary
# Write any fractions as fractions (i.e. p/q), not decimals
a, b, c, d = sp.symbols('a b c d')

sol5 = ...
```

```python
grader.check("q5b")
```

#### 🔴   Question 5e (2 pts) ###


Rather than looking for a degree three polynomial, suppose we wanted to find a polynomial that passes through the four points and that has degree two, such as

$$p(x) = a + bx+cx^2$$

i). Write the augmented matrix for the four linear equations for the coefficients obtained by requiring that the graph of this 2nd degree polynomial passes through the four points above.

ii). Then find the RREF for the augmented matrix.

iii).   Based on the RREF matrix you found above, choose the answer that best describes the system of linear equations and then assign the the variable  `ans_5e` to the corresponding letter (as a capital letter string `'A'`,`'B'`, or `'C'`) to  below to indicate your answer.  Replace the ellipses (`...`) with your answer. 

     A).  The system is inconsistent

     B).  The system is consistent and the solution is unique

     C).  The system is consistent and the solution is not unique.  
                                                                               
                                                            
 

```python
#(i)
q5e_aug = ...

#(ii)
q5e_rref= ...

q5e_rref
```

```python
#(iii)
ans_5e = ...
```

```python
grader.check("q5e")
```

#### 🔴   Question 5f (5 pts) ###

Rather than looking for a degree 3 polynomial, suppose we wanted to find a polynomial that passes through the four points and that has degree 4, such as

$$p(x) = a + bx+cx^2+dx^3+ex^4$$

i). Write the augmented matrix for the four linear equations for the coefficients obtained by requiring that the graph of this 4th degree polynomial passes through the four points above.

ii). Then find the RREF for the augmented matrix.

                  
iii).   Based on the RREF matrix you found above, choose the answer that best describes the system of linear equations and then assign the the variable  `ans_5f` to the corresponding letter (as a capital letter string `'A'`,`'B'`, or `'C'`) to  below to indicate your answer.  Replace the ellipses (`...`) with your answer. 

     A).  The system is inconsistent

     B).  The system is consistent and the solution is unique

     C).  The system is consistent and the solution is not unique.  
                                                                               
                                                                                 
iv). Use the RREF to find the solution set to the original system of equations.   Do not rename any free variables.
```python

a, b, c, d , e = sp.symbols('a b c d e')
sol5f = {
    a: answer(s), 
    b: answer(s),
    c: answer(s)
    d: answer(s)
    e: answer(s)
}

 ```
If no solution exists, leave the dictionary empty, that is sol = { } 

Check your work by graphing the polynomial(s) you found in Desmos and verifying that it passes through the four points.   https://www.desmos.com/calculator


```python

#(i)
q5f_aug = ...


#(ii)
q5f_rref = ...


```

```python
#(iii)
ans_5f = ...
```

```python
#(iv)

# Use the RREF to find the solution set to the system of linear equations.  Write your answer as a Sympy dictionary
# Write any fractions as fractions (i.e. p/q), not decimals
a, b, c, d, e = sp.symbols('a b c d e')

sol5f = ...
```

```python
grader.check("q5f")
```

#### 🔴   Question 5g (1 pt) ###

Suppose you had 10 points $(x,y)$ (with 10 unique $x$ values) and you wanted to find a polynomial passing through each of them. What should the degree of the polynomial be to guarantee that there is exactly one such polynomial?

```python
deg = ...
```

```python
grader.check("q5g")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴     Question 6  ###

Linear systems have applications in chemistry when balancing chemical equations. When a chemical reaction occurs, molecules of different substances combine to create molecules of other substances. Chemists represent such reactions with chemical equations. To balance a chemical equation means to find the number of atoms of each element involved that will preserve the number of

atoms in the reaction. 

As an example, consider the chemical equation

$$C_2H_6 + O_2 → CO_2 + H_2O$$


This equation asks about what will happen when the chemicals ethane ($C_2H_6$) and oxygen ($O_2$),
called the reactants of the reaction, combine to produce carbon dioxide ($CO_2$) and water ($H_2O$),
called the products of the reaction (note that oxygen gas is diatomic, so that oxygen atoms are
paired). 


The arrow indicates that it is the reactants that combine to form the products. 

Any chemical reaction has to obey the Law of Conservation of Mass that says that mass can neither be created
nor destroyed in a chemical reaction.  Consequently, a chemical reaction requires the same number
of atoms on both sides of the reaction. 

In other words, the total mass of the reactants must equal
the total mass of the products. In the reaction above, the chemicals involved are made up of carbon
(C), hydrogen (H), and oxygen (O) atoms. To balance the equation, we need to know how many
molecules of each chemical are combined to preserve the number of atoms of C, H, and O.



Let 
 - $x_1$ be the number of molecules of $C_2H_6$, 
 - $x_2$ the number of molecules of $O_2$, 
 - $x_3$ the number of molecules of $CO_2$, 
 - $x_4$ the number of molecules of $H_2O$ in the reaction. 

We can then represent this reaction as

$$x_1C_2H_6 + x_2O_2 \to x_3CO_2+x_4H_2O$$

In each molecule (e.g., ethane $C_2H_6$), the subscripts indicate the number of atoms of each
element in the molecule. So 1 molecule of ethane contains 2 atoms of carbon and 6 atoms of
hydrogen. Thus, there are 2 atoms of carbon in $C_2H_6$ and 0 atoms of carbon in $O_2$, giving us
$2x_1$ carbon atoms in $x_1$ molecules of $C_2H_6$ and 0 carbon atoms in $x_2$ molecules of $O_2$. 

On the
product side of the reaction there is 1 carbon atom in $CO_2$ and 0 carbon atoms in $H_2O$. To balance
the reaction, we know that the number of carbon atoms in the products must equal the number of
carbon atoms in the reactants.

<!-- BEGIN QUESTION -->

#### 🔴     Question 6a (2 pt) ###
Write the system of linear equations that models this chemical reaction (hint: set up an equation that balances the number of carbon atoms on both sides of the reaction, then repeat with hydrogen and oxygen). 

#### Question 6a Solution) 

**Type your work answering Question 6a** in this cell, using the align format shown above.    Do not add any additional cells to this part.  

<!-- END QUESTION -->

## 🔴      Question 6b (5 pts) ###

i). Write the augmented matrix for this system of equations.

ii). Then find the RREF for the augmented matrix.

                  
iii).   Based on the RREF matrix you found above, choose the answer that best describes the system of linear equations and then assign the the variable  `ans_6` to the corresponding letter (as a capital letter string `'A'`,`'B'`, or `'C'`) to  below to indicate your answer.  Replace the ellipses (`...`) with your answer. 

     A).  The system is inconsistent

     B).  The system is consistent and the solution is unique

     C).  The system is consistent and the solution is not unique.  
                                                                               
                                                                                 
iv). Use the RREF to find the solution set to the original system of equations.   Do not rename any free variables.
```python

x1, x2, x3, x4  = sp.symbols('x1 x2 x3 x4')
sol6 = {
    x1: answer(s), 
    x2: answer(s),
    x3: answer(s)
    x4: answer(s)
}

 ```
If no solution exists, leave the dictionary empty, that is sol = { } 


```python

#(i)
q6_aug = ...

#(ii)
q6_rref = ...



q6_rref
```

```python
#(iii)

ans_6 = 'C'
```

```python
# (iv)

# Use the RREF to find the solution set to the system of linear equations.  Write your answer as a Sympy dictionary
# Write any fractions as fractions (i.e. p/q), not decimals
x1, x2, x3, x4 = sp.symbols('x1, x2, x3, x4')

sol6 = ...


```

```python
grader.check("q6b")
```

## <span style='color:Red'>   Question 6c (2 pts) ###

Note that we cannot have a fraction of a molecule in our reaction.  To "balance" the reaction (and get a unique solution), we find the minimum values for $x_1$, $x_2$, $x_3$ and $x_4$ that don't have fractions of molecules in the reaction.  
                                                                                                                         
                                                                                                                    

```python
# Enter the balanced reaction solution here:


x_1 = ...

x_2 = ...


x_3 = ...

x_4 = ...
```

```python
grader.check("q6c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 7  ###



### Operation Counts for RREF
Let's examine the computational costs associated with RREF.  Keep in mind though that raw computational cost is not the only thing that determines the performance of an algorithm.  In future lessons we'll look at how memory access of an algorithm can play a large role in the overall performance of an implementation.  


On typical modern architectures there is only a slight difference between the time it takes to multiply or add two floating-point numbers, so we'll lump these two operations together and refer to them as FLOPS (floating-point operations). 

**Getting matrices into RREF**: Assume that $A$ is a real $n \times n$ matrix given by 


 $A=\left[{ 
\begin{align}
   a_{11} \hspace{2mm} & a_{12} \hspace{2mm}& & \cdots & & a_{1n} \\
a_{21}\hspace{2mm} & a_{22} & & \cdots & & a_{2n} \\
\vdots \hspace{2mm} & \vdots & &        & & \vdots \\
a_{m1}\hspace{2mm} & a_{m2} & & \cdots & & a_{mn} \\
  \end{align}  }\right]$






Our goal is to transform a matrix to its reduced row echelon
form, which looks something like this:

 $\left[{ 
\begin{align}
1 &\hspace{2mm} 0 &\hspace{2mm}0  & \cdots &0 &\hspace{2mm} 0 \\
0 & \hspace{2mm}1 &\hspace{2mm} 0 & \cdots &0 &\hspace{2mm} 0 \\
\vdots &\hspace{2mm} \vdots & &        & & \vdots \\
0 &\hspace{2mm} 0 &\hspace{2mm} & \cdots & & \hspace{2mm}1 \\
 \end{align}  
}\right]$



 We are mainly interested in the case when $n$ 
 is very large, which is when we need to worry about how much effort is required.

Let’s first consider the effort required to replace one entry in the matrix with a 0.    

                 


## <span style='color:Red'>   Question 7a (2 ps) ###

**Multiple Choice:**
Given an $n$x$n$ matrix, how many FLOPS does it take to replace one entry in the matrix with a $0$ (by taking the row and replacing it with a multiple of another row)?

```python
n = sp.symbols('n')

ans_7a = ...


```

```python
grader.check("q7a")
```

We perform one replacement operation for every 0 entry
	in the reduced row echelon matrix.  

When $n$ is very
	large, most of the $n^2$ entries in the reduced row
	echelon form are 0 so we need roughly $n^2$
	replacements.  


Thus the number of FLOPS
	resulting from the needed replacements is roughly $n^2(2n)$ =
	$2n^3$
    
    
   To make the first non-zero term in each row a 1, each row is scaled roughly one time so there are roughly
    $n$ scaling operations, each of which requires $n$
    operations.  The number of FLOPS due to scaling is roughly
    $n^2$.

Therefore, the total number of FLOPS to get a matrix into RREF is roughly
     $$2n^3 + n^2$$.  






When $n$ is very large, the
      $n^2$ term is much smaller than the $n^3$ term. 

(Caveat: This is a very rough measure of the effort required to find the reduced row echelon form; **a more careful accounting shows that the number of FLOPS is roughly $\frac{2}{3}n^3$**).  

As we have seen, some matrices require more effort than others, but the upshot of this observation is that the effort is proportional to $n^3.$

We can think of this in the following way: If the size of the matrix grows by a factor of 10, then the effort required grows by a factor of $10^3$.





## <span style='color:Red'>   Question 7b (2 pts) ###




 To put this into context, imagine we need to solve a linear
      system with 10 billion equations and 10 billion variables.  Suppose we have a computer that can perform a trillion
       operations every second. 
How long (in years) would it take to find the RREF for this system?  (Assume 365 days in a year and **use the more careful accounting of FLOPS mentioned above**).  


```python
flops = ...

int(flops)
```

```python
grader.check("q7b")
```



This may seem like an absurd situation, but we use the results
      of such a computation **every day.**.
      

The matrices used in Google's PageRank algorithm are extremely large, with dimensions corresponding to the number of web pages indexed by Google. This number can reach into the billions.
Specifically, the "Google matrix" or "connectivity matrix" represents the entire web graph, where each row and column corresponds to a web page, and a non-zero entry indicates a hyperlink between pages. Given the vastness of the internet, this results in matrices with dimensions potentially exceeding ten billion by ten billion.



Clearly, we will need some better tools to deal with
      really big problems like this one - we will soon learn these in our course!

	
      
     
      

     



          

  

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## <span style='color:Red'>   Question 8  ###



In Question 7, we looked at one computational limitation of solving systems of equations this way: once a matrix gets to be too big, it is not reasonable to apply Gaussian elimination to find its reduced row echelon form.
    
In this exercise, we will see another limitation: computer arithmetic with real numbers is only an approximation because computers represent real numbers with only a finite number of bits. For instance, the number pi:

$$\pi =3.141592653589793238462643383279502884197169399\ldots $$

would be approximated inside a computer by, say,

$$\pi\approx 3.141592653589793$$

Most of the time, this is not a problem. However, when we perform millions or even billions of arithmetic operations, the error in these approximations starts to accumulate and can lead to results that are wildly inaccurate. 

Here are two examples demonstrating this.

## <span style='color:Red'>   Question 8a (1 pt) ###

Let’s first see an example showing that computer arithmetic really is an approximation. First, consider the linear system


\begin{align} 
x+\frac{1}{2}y+\frac{1}{3}z &= 1 \\ 
\frac{1}{2}x + \frac{1}{3}y+\frac{1}{4}z &= 0 \\
\frac{1}{3}x+\frac{1}{4}y+ \frac{1}{5}z &= 0
\end{align}



Use SymPy to write the system as an augmented matrix and then find the RREF.  

```python

q8a_aug = ...

q8a_rref = ...

q8a_rref
```

```python
grader.check("q8a")
```

## <span style='color:Red'>   Question 8b (1 pt) ###


Notice there are floating point errors in the RREF matrix above.  What is the correct solution to this system of equations (without floating point errors)?


```python

x = ...
y = ...
z = ...

```

```python
grader.check("q8b")
```

Fortunately, since SymPy is a symbolic system, it can use symbolic maninpulation to find the RREF for this system and thus give an answer without floating point errors.  However, first you need to define each fraction using the SymPy method `Rational` as follows:


```python
sp.Rational(1,2)
```

## <span style='color:Red'>   Question 8c (1 pt) ###




Use SymPy to rewrite the system as an augmented matrix where any fractions are rewritten using the `sp.Rational` function then find the RREF.  

```python

q8c_aug = ...

q8c_rref = ...

q8c_rref
```

```python
grader.check("q8c")
```

## <span style='color:Red'>   Question 8d (1 pt) ###

Another problem that arises is some types of linear systems are particularly sensitive to errors resulting from computers’ approximate arithmetic. For instance, suppose we are interested in the linear system

\begin{align} 
x+y &= 2 \\ 
x+1.001y &= 2 \\ 
\end{align}



Use SymPy to write the system as an augmented matrix, find the RREF and find the solution to the system.  

```python

q8d_aug = ...

q8d_rref = q8d_aug.rref(pivots = False) #

q8d_rref
```

```python
#Solution to system:

q8d_x = ...
q8d_y = ...
```

```python
grader.check("q8d")
```

## <span style='color:Red'>   Question 8e (1 pt) ###

Suppose now that the computer has accumulated some error in one of the entries of this system so that it incorrectly stores the system as

\begin{align} 
x+y &= 2 \\ 
x+1.001y &= 2.001 \\ 
\end{align}



Use SymPy to write the system as an augmented matrix, find the RREF and find the solution to the system.  

```python

q8e_aug = ...

q8e_rref = ...

q8e_rref
```

```python
#Solution to system:

q8e_x = ...
q8e_y = ...
```

```python
grader.check("q8e")
```

Notice how a small error in one of the entries in the linear system leads to a solution that has a dramatically large error!  

We call this type of linear system **ill-conditioned**. 

Fortunately, this is an issue that has been well studied, and there are techniques that we will learn later in this class that mitigate this type of behavior.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


### Submission Instructions

Before proceeding any further, **save this notebook.**

Then run the cell below to double check that you don't have any spaces between dollar signs and text when writing LaTeX:

```python
# Run this cell before you run the 'grader.export()' cell below.  
# It will search for LaTeX errors that will cause the LaTeX compiler to fail.  

import simple_latex_checker as slc

nb = slc.Nb_checker()
nb.run_check("lab01.ipynb")
```

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 


* **Lab 1 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the Lab 1 Autograded assignment in Gradescope.

* **Lab 1 Manually Graded**: Submit your lab01.PDF to the Lab 1 Manually Graded assignment in Gradescope.  **It is your responsibility to fully review your PDF file before submitting and make sure that all your lines of code are visible and any LaTeX has correctly compiled and is fully viewable.**  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.** If not, you will lose points.    

[TROUBLESHOOTING TIPS](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing)  If you are having any issues compiling your assignment, please read through these troubleshooting tips first, then post any questions on Piazza.  

**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='lab01.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1ab (points: 2)
```
>>> assert hash_simplified_matrix(q1_aug) == '940e56cc85da0abf9409025bec5d1730441d09f1fe42f2c581e84fe210310c9d'
```
```
>>> assert hash_simplified_matrix(q1_rref) == '6c4e896d7e9812cc5c886f8c4777f6b60c964589187588b4cf6cedb951bcd155'
```


Tests for q1c (points: 1)
```
>>> assert get_hash(ans_1c) == '285c343c9a3026f47671e8f8ea859d92'
```


Tests for q1d (points: 1)
```
>>> assert isinstance(sol1, dict), 'sol1 must be a dictionary'
```
```
>>> assert set(sol1.keys()) == {x, y, z}, 'keys must be x, y and z'
```
```
>>> assert get_hash(sol1[x]) == '22d38fd62431fc64b9a7aff472b11dfe'
```
```
>>> assert get_hash(sol1[y]) == '2207903ae69f4087640740088c780f15'
```
```
>>> assert get_hash(sol1[z]) == 'd0a08be6db077fee6d9bdcbf75cec5c2'
```


Tests for q2 (points: 4)
```
>>> assert hash_simplified_matrix(q2_aug) == '6ecfbb01178f16d7199a8c04d8508a9b10d4427f74b8e9c33ecde70d33d55bea'
```
```
>>> assert hash_simplified_matrix(q2_rref) == '380ce0730964df89fdc45817c2321c689ab1ca4751b04356ba85db9a2b3bf3bd'
```
```
>>> assert get_hash(ans_2c) == '0078a48aea0a05c2f49c537dda19713f'
```
```
>>> assert isinstance(sol2, dict)
```
```
>>> assert set(sol2.keys()) == {x1, x2, x3, x4, x5}
```
```
>>> assert get_hash_symb(sol2[x1]) == '3416cc2a3b51ea8e6f2b126c5f93f3bf'
```
```
>>> assert get_hash_symb(sol2[x2]) == 'd678d0b2e5eb10187dd765801619471e'
```
```
>>> assert get_hash_symb(sol2[x3]) == '5201a4a5e1caeb5c7525d66fb5519255'
```
```
>>> assert get_hash_symb(sol2[x4]) == '428fce595c9adf6e7e4f85fd63728a00'
```
```
>>> assert get_hash_symb(sol2[x5]) == '22d38fd62431fc64b9a7aff472b11dfe'
```


Tests for q3b (points: 3)
```
>>> assert isinstance(sol3, dict)
```
```
>>> assert set(sol3.keys()) == {m1, m2, m3}
```
```
>>> assert get_hash(sol3[m1]) == '5c81e6c563fd762aa608c02c30cb3324'
```
```
>>> assert get_hash(sol3[m2]) == 'eda28df8259ee4793c70291bafcd6186'
```
```
>>> assert get_hash(sol3[m3]) == 'c4c1f6625a2ee8ceb8b5d67a1b6840b8'
```


Tests for q4b (points: 4)
```
>>> assert hash_simplified_matrix(q4_rref) == 'e9d147f5fe4bf0bb9e75890bdda34a85b820858325e89cfee160e43f689588ac'
```
```
>>> assert get_hash(ans_4d) == '0078a48aea0a05c2f49c537dda19713f'
```
```
>>> assert isinstance(sol4, dict)
```
```
>>> assert set(sol4.keys()) == {w, x, y, z}
```
```
>>> assert get_hash_symb(sol4[w]) == '8aff42df53cddbb9518d280c765e6be2'
```
```
>>> assert get_hash_symb(sol4[x]) == '693c163db98fe8782bd3421346c95917'
```
```
>>> assert get_hash_symb(sol4[y]) == '732fb38446129b4f4cbf499239240ccf'
```
```
>>> assert get_hash_symb(sol4[z]) == '5179901de67cd258090e45007346eac5'
```


Tests for q5b (points: 4)
```
>>> assert hash_simplified_matrix(q5_rref) == '4d538928f12ced411c8234144c2d8b01c2c2edb13784931c417e0950866815a0'
```
```
>>> assert get_hash(ans_5d) == '285c343c9a3026f47671e8f8ea859d92'
```
```
>>> assert isinstance(sol5, dict)
```
```
>>> assert set(sol5.keys()) == {a, b, c, d}
```
```
>>> assert get_hash_symb(sol5[a]) == '6519cc55d84e853e0e3206be496ebf51'
```
```
>>> assert get_hash_symb(sol5[b]) == 'dee977ea1fbfa26442fcf1a8f361ce7f'
```
```
>>> assert get_hash_symb(sol5[c]) == '179b5d1eca09792e9ce38f283fbcf52e'
```
```
>>> assert get_hash_symb(sol5[d]) == '75dab9a87e25210f8c4bfb42bdd2f112'
```


Tests for q5e (points: 2)
```
>>> assert hash_simplified_matrix(q5e_rref) == '4dd5c7c7a901c154ecf89f3bb7ae3f55ef21dd10a768a167ab6cfaebe7ca24e6'
```
```
>>> assert get_hash(ans_5e) == '44614839b080dfac0c3a41265ccf0ace'
```


Tests for q5f (points: 5)
```
>>> assert hash_simplified_matrix(q5f_rref) == 'a0956d7031307fac1b8ed3c24ce39c85ebc3f256c1a734e829ed59ad8b8ae2d8'
```
```
>>> assert get_hash(ans_5f) == '0078a48aea0a05c2f49c537dda19713f'
```
```
>>> assert isinstance(sol5f, dict)
```
```
>>> assert set(sol5f.keys()) == {a, b, c, d, e}
```
```
>>> assert get_hash_symb(sol5f[a]) == '6519cc55d84e853e0e3206be496ebf51'
```
```
>>> assert get_hash_symb(sol5f[b]) == '9231477f0016f836ad2870de43868ec5'
```
```
>>> assert get_hash_symb(sol5f[c]) == '040a332cbec4d508b9d625501e361aa4'
```
```
>>> assert get_hash_symb(sol5f[d]) == '64014081c00f6e54b6097df2ef2d900e'
```
```
>>> assert get_hash_symb(sol5f[e]) == 'c2fad9fedd6b5c3f027a65edef8cd7eb'
```


Tests for q5g (points: 1)
```
>>> assert get_hash(deg) == '085fdd39dfb189c7650225c26b4c2fbc'
```


Tests for q6b (points: 5)
```
>>> assert hash_simplified_matrix(q6_rref) == 'a9709a7c07031dad4b8cab24956f4d119bac1a29fbd8f61942101fc680568991'
```
```
>>> assert get_hash(ans_6) == '0078a48aea0a05c2f49c537dda19713f'
```
```
>>> assert isinstance(sol6, dict)
```
```
>>> assert set(sol6.keys()) == {x1, x2, x3, x4}
```
```
>>> assert get_hash_symb(sol6[x1]) == 'bc50268f5f36fa2be81b723f238fbc24'
```
```
>>> assert get_hash_symb(sol6[x2]) == '5643de434db153f750ce366bb19fd9a8'
```
```
>>> assert get_hash_symb(sol6[x3]) == '26971159ff5b52b3cb071a11f984fc54'
```
```
>>> assert get_hash_symb(sol6[x4]) == '428fce595c9adf6e7e4f85fd63728a00'
```


Tests for q6c (points: 2)
```
>>> assert get_hash(x_1) == '6519cc55d84e853e0e3206be496ebf51'
```
```
>>> assert get_hash(x_2) == 'a608ec14f8da402c79175fa26695777e'
```
```
>>> assert get_hash(x_3) == 'ad8651940e2ebabeb839530d201156b7'
```
```
>>> assert get_hash(x_4) == '5c81e6c563fd762aa608c02c30cb3324'
```


Tests for q7a (points: 2)
```
>>> assert get_hash(ans_7a) == '8018c2aa5af6e3c925fbc3e5a7a25206'
```


Tests for q7b (points: 2)
```
>>> assert get_hash(int(flops)) == '09074a552401cb53e3a6777743ef78ad'
```


Tests for q8a (points: 1)
```
>>> assert hash_simplified_matrix_floats(q8a_rref) == '1412fe8ee6adeaad0956ab686829e2ccdc499ab2cba2167088c3699b70739bb4'
```


Tests for q8b (points: 1)
```
>>> assert get_hash(x) == '085fdd39dfb189c7650225c26b4c2fbc'
```
```
>>> assert get_hash(y) == '2a571686221be713519532be035ba686'
```
```
>>> assert get_hash(z) == 'cfc3bd433703abc7938aa522f5a2aaec'
```


Tests for q8c (points: 1)
```
>>> assert hash_simplified_matrix(q8c_rref) == '1412fe8ee6adeaad0956ab686829e2ccdc499ab2cba2167088c3699b70739bb4'
```


Tests for q8d (points: 1)
```
>>> assert get_hash(q8d_x) == '6519cc55d84e853e0e3206be496ebf51'
```
```
>>> assert get_hash(q8d_y) == 'd0a08be6db077fee6d9bdcbf75cec5c2'
```


Tests for q8e (points: 1)
```
>>> assert get_hash(q8e_x) == '7fc7dddc3020d697a94d29c4eb434309'
```
```
>>> assert get_hash(q8e_y) == '7fc7dddc3020d697a94d29c4eb434309'
```
